# Example Task

This is an example task. Students are asked to write code that produces some output.

## Manually graded

The below is to be manually graded. First, a task ("manually graded task").

Do something cool!

Next, an answer ("manually graded answer"):

<!-- grade_id: cell-61bcab10cedd4aa4, points: 2, grade: true, solution: true -->
```python
"""This is an incorrect answer."""
```

## Autograded
There are also autograded answers, with tests. The task here is to assign `a=50`.

<!-- grade_id: cell-55e352987e3b2fbf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
a = 20
```

The tests are read-only:

<!-- grade_id: cell-d81b5763f9a533ab, points: 2, grade: true, solution: false -->
```python
assert a == 50
```